# EEN065 - Applied object-oriented programming


Teacher: [Seyed Morteza Ahmadian](https://www.chalmers.se/en/persons/seyedah/) and [Carlos Natalino](https://www.chalmers.se/en/persons/carda/) / Examiner: [Paolo Monti](https://www.chalmers.se/en/persons/mpaolo/)

[Canvas course page](https://chalmers.instructure.com/courses/39015)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before submitting this assignment, please ensure that everything runs correctly.

First, **restart the kernel**, then **run all cells**.
Next, verify that all tests execute successfully.

Note that if any problem produces an error, the subsequent problems **will not** be tested.

In case of discrepancies between the problem description and the tests, you should resolve the issue with the tests in mind, as they determine correctness.

There are four types of cells in this notebook:
1.	*Solution cells*: These are the cells where you write your answers or modify the existing code to solve the problem.
2.	*Debug cells*: These cells are for debugging your solution. You may add or edit code here, but ensure that it does not introduce any errors.
3. *Test cells:*  These cells run **automated tests** to verify whether your solution is **functionally correct**.
4. *Code quality tests:*  These cells check whether your solution follows the **code quality standards** of the course. Even if your code produces the correct output, it must also be **clean, readable, and well-structured**.

- If all tests pass, you will see an output similar to:
  - `1 passed`

- If a test fails, you will see an **error message** explaining:
  - Which input was tested
  - What output was expected
  - What your function actually returned

You will see several sections in the output:

- **tests:** Runs the test cells with coverage enabled.  
  - `1 passed` indicates success.
- **coverage:** Checks for any line of your solution that is not useful.  
  - `100% coverage` indicates success.
- **flake8:** Checks for style issues (e.g., unused variables, line length).  
  - Any output other than `0` here means you must fix style violations.
- **pylint:** Performs deeper code analysis and assigns a score.  
  - A score of `10.00/10` means your code fully complies.
- **mypy**: Checks that your code respects the specified type annotations.
  - Any reported error means a type hint is missing or incorrect.
- **isort:** Verifies that imports are correctly ordered.  
  - If differences are shown, reorder your imports.

At the end:
- **`tests passed ✓`** → Your solution is correct *and* meets all quality requirements.
- **`your code presented some issue`** → Review the messages above, fix the reported issues, and rerun the cell.

**You must pass both functional tests and code quality checks to receive full credit.**

Common errors and what they mean:
- **`You did not execute your solution cell!`**  
  → You forgot to run the solution cell before running the test cell.
- **`The module does not have the necessary function!`**  
  → The required function is missing or named incorrectly.
- **AssertionError**  
  → Your function ran, but it returned an incorrect result for at least one test case.

**Fix all test errors before moving on**, as later problems will not be tested if an earlier one fails.

**Delete** the line `raise NotImplementedError()` from any problem you solve.

**Do not add or remove any cells in this file.** All required cells are already provided.

To execute a cell, select it and press **CTRL+Enter** (Windows), **CMD+Enter** (macOS), or click the **Run cell** button.

## Assignment Week 6

**Observations:**
For the problems in this assignment, you should:
- make sure that your `utils.py` file has the following line.
- you will find the set of references needed to solve this assignment in the assignment page in canvas.

**Preparation:** Run the cell below every time you start working on this file, and every time you restart the kernel.

```python
%load_ext autoreload
import os

# getting latest version of utils.py
# import requests
# response = requests.get("https://raw.githubusercontent.com/carlosnatalino/programming-assignments/refs/heads/main/utils.py")
# with open("utils.py", "wb") as f:
#     f.write(response.content)

if not os.path.exists("utils.py"):
    raise ValueError("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.")

# check if this file is within the same folder as the `utils.py` file.
try:
    from utils import export_cell_content
except:
    import sys
    raise ValueError("Did you select the right kernel that starts with venv?")

filename = "assignment_week_06_public.ipynb"
assert filename != None, "The file name must be set."
if not os.path.exists(filename):
    raise ValueError(f"The value of the `filename` variable is wrong. \
The file `{filename}` does not exist. Write the correct name of this file!")

os.environ["NOTEBOOK_NAME"] = os.path.join(os.getcwd(), filename)
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

### Task 1

Create a Python class with the name `VolumeCalculator`.

The initialization `.__init__()` method receives a number and assign the value to the attribute `radius`.

Then, create a method named `volume()` that returns the volume of the sphere using the `radius` attribute.
If the `radius` attribute is `None` or $\leq 0$ the method `volume()` should return `None`.

**Remember:** from the `math` module, import `pi`.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# solution cell
from math import pi

class VolumeCalculator:
    
    def __init__(self, radius):
        self.radius = radius

    def volume(self):
        if self.radius is None or self.radius <= 0:
            return None
        
        return (4/3) * pi * (self.radius ** 3)
```

```python
# uncomment the line(s) below to debug

radius = 2
s = VolumeCalculator(radius)
print(f'The volume of a sphere with radius {radius} is {s.volume()}')
print('Execution finished', u'\u2713')
```

Output:
```
The volume of a sphere with radius 2 is 33.510321638291124
Execution finished ✓
```

```python
# test cell
def test_solution_1() -> None:
    try:
        import solution_1
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_1 import VolumeCalculator
    except:
        raise ValueError("The class does not exist. \
            Did you execute your solution cell?")

    assert 'volume' in dir(VolumeCalculator), 'Method volume() does not exist!'

    test_cases = [(None, None), (0, None), (-1, None), (1, 4.1887902047863905), \
        (2, 33.510321638291124)]
    for _in, _out in test_cases:
        _res = VolumeCalculator(_in)
        assert _res.volume() == _out, f"The volume function with input `{_in}` should \
            return the value `{_out}` of type `{type(_out)}`\n \
                but returned the value `{_res}` of type `{type(_res)}`."
```

```python
# validate the solution
solution_file = "solution_1.py"
test_file = "test_solution_1.py"
export_cell_content("solution_1", solution_file)
export_cell_content("tests_1", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    print('\033[91m your code presented some issue, check messages above', u'\u274C', \
        "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-05-08 10:06:46.475085 (2 seconds ago)
Latest saved date of the file: 2026-05-08 10:06:46.475085 (2 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[0m
############################################# coverage #############################################

[32m[32m[1m1 passed[0m[32m in 0.04s[0m[0m
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_1.py       8      0      2      0   100%
-----------------------------------------------------------
TOTAL               8      0      2      0   100%
coverage True

[0m
############################################## flake8 ##############################################
solution_1.py:4:1: E302 expected 2 blank lines, found 1
class VolumeCalculator:
^
solution
```

### Task 2

Create a Python class named `Person`.
In the initialization method, the parameters are:
1. name (mandatory)
2. role (mandatory)
3. person number (mandatory)
4. address (optional, `None` by default)

The initialization method should save these parameters as attributes.

The class must also implement `__repr__()` and return a string containing only the attributes that are not None, separated by `\t` characters.

Example: "John Doe\tManager\t201414202014".

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# solution cell
raise ValueError("Delete this line and write your solution here!")
```

```python
# uncomment the line(s) below to debug

person = Person('John Doe', 'Manager', '201414202014')
print(f'The Person object converts to string as {person}')
print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
# test cell
from unittest.mock import patch


def test_solution_2() -> None:
    try:
        import solution_2
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_2 import Person
    except:
        raise ValueError("The class does not exist. \
            Did you execute your solution cell?")

    p = Person('John Doe', 'Manager', '201414202014')

    with patch('builtins.print') as mock_print:
        s = str(p)
    mock_print.assert_not_called()

    test_cases = [(('John Doe', 'Manager', '201414202014'), \
        'John Doe\tManager\t201414202014'), \
        (('Jane Doe', 'Supervisor', '201212202012', 'Gibraltargatan'), \
            'Jane Doe\tSupervisor\t201212202012\tGibraltargatan'), \
        (('Jane Doe', 'Supervisor', '201212202012', 'Gibraltargatan'), \
        'Jane Doe\tSupervisor\t201212202012\tGibraltargatan')]
    for _in, _out in test_cases:
        _res = Person(*_in)
        assert repr(_res) == _out, f"The Person class with input `{_in}` should be \
            displayed as the value `{_out}` of type `{type(_out)}`\n \
                but returned the value `{_res}` of type `{type(_res)}`."
```

```python
# validate the solution
solution_file = "solution_2.py"
test_file = "test_solution_2.py"
export_cell_content("solution_2", solution_file)
export_cell_content("tests_2", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    print('\033[91m your code presented some issue, check messages above', u'\u274C', \
        "\033[0m")
```

### Task 3

Create a Python dataclass named `PersonAge`.
The attributes of this class are:
1. name (Type: string)
2. job (Type: string)
3. birthdate (Type: datetime.date)
4. age (Type: int, not given as an argument during initialization)

The attributes `name`, `job` and `birthdate` are mandatory during initialization.
Then, you need to create a method `.compute_age()` that uses the `birthdate` to compute how old the person is, saves it to the `age` property, and returns the `age` value.

Remember to import `dataclass` and `field` from `dataclasses` module.
Also import `datetime` for birthdate attribute.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# solution cell
raise ValueError("Delete this line and write your solution here!")
```

```python
# uncomment the line(s) below to debug
from datetime import date

person1 = PersonAge('John Doe', 'Manager', date(1990, 11, 15))
person2 = PersonAge('Jane Doe', 'Manager', date(1980, 11, 15))
print(f"The age of {person1.name} is {person1.compute_age()}")
print(f"The age of {person2.name} is {person2.compute_age()}")
print('Execution finished', u'\u2713')
```

```python
# test cell
import dataclasses
from datetime import date
from unittest.mock import patch

def test_solution_3() -> None:
    try:
        import solution_3
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_3 import PersonAge
    except:
        raise ValueError("The class does not exist. Did you execute your solution cell?")

    assert dataclasses.is_dataclass(PersonAge), "The class is not a dataclass!"
    assert any(f.name == "age" and f.init is False for f in dataclasses.fields(PersonAge))
    
    p = PersonAge('John Doe', 'Manager', date(1990, 11, 15))
    assert hasattr(p, 'name'), 'Class PersonAge does not have property `name`.'
    assert hasattr(p, 'job'), 'Class PersonAge does not have property `job`.'
    assert hasattr(p, 'age') == False, 'Class PersonAge should not have property `age`.'
    p.compute_age()
    assert hasattr(p, 'age'), 'Class PersonAge does not have property `age`.'
    assert p.age == 35, f"The computed age `{p.age}` is incorrect!"

    with patch('__main__.print') as mock_print:
        s = str(p)
    mock_print.assert_not_called()

    person1 = PersonAge('John Doe', 'Manager', date(1990, 11, 15))
    person1.compute_age()
    person2 = PersonAge('Jane Doe', 'Manager', date(1980, 11, 15))
    person2.compute_age()
    assert person1.age < person2.age, "The age computation seems incorrect!"
```

```python
# validate the solution
solution_file = "solution_3.py"
test_file = "test_solution_3.py"
export_cell_content("solution_3", solution_file)
export_cell_content("tests_3", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    print('\033[91m your code presented some issue, check messages above', u'\u274C', \
        "\033[0m")
```

### Task 4

Create a Python class named `Student` that inherits from the class `Person` created in task 2. The class `Student` should override the initialization method of the class `Person` and set the attribute `role` always with the value `"Student"`, removing the need to pass it as a parameter to the method. The initialization method should also include the parameter `program` as a mandatory one, before the phone_number and address. That is, you need to create an initialization method with parameters:
1. name (mandatory)
2. person number (mandatory)
3. program (mandatory)
5. address (optional, `None` by default)

Finally, the class `Student` should also have a method `.print()` that uses the return of the method `.__repr__()` of the *parent* class and appends program to it, printing out the result.

**Suggestions:**
- You must use `super().__init__()` to initialize the attributes inherited from the class `Person`.
- Since the class in this problem inherits from the class created in the previous task, note that you need to import the `Person` class from the module created in the previous task in the solution cell of this one.
- To call a method from the parent class you need to use the function `super()`. For example, to call the `.__init__()` method from the parent class, you call `super().__init__()`.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# solution cell
raise ValueError("Delete this line and write your solution here!")
```

```python
# uncomment the line(s) below to debug

s = Student('John Doe', '201414202014', 'TIDSL')
print('Student:')
s.print()

s = Student('John Doe', '201414202014', 'TKDES', 'Gibraltargatan')
print('Student:')
s.print()
print('Execution finished', u'\u2713')
```

```python
# test cell
from unittest.mock import patch


def test_solution_4() -> None:
    try:
        import solution_4
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_4 import Student
        from solution_2 import Person
    except:
        raise ValueError("Your solution does not contain the right function!")

    assert 'print' in dir(Student), 'Class Student does not have a method `print`.'

    s = Student('John Doe', '201414202014', 'TIDSL')
    assert isinstance(s, Person), 'The class Student is not a child class of the class Person'
    with patch('solution_4.print') as mock_print:
        s.print()
    mock_print.assert_called_once_with('John Doe\tStudent\t201414202014\tTIDSL')

    s = Student('John Doe', '201414202014', 'TKDES', 'Gibraltargatan')
    with patch('solution_4.print') as mock_print:
        s.print()
    mock_print.assert_called_once_with('John Doe\tStudent\t201414202014\tGibraltargatan\tTKDES')
    assert hasattr(s, 'program'), 'The object does not seem to have the property `program`'
```

```python
# validate the solution
solution_file = "solution_4.py"
test_file = "test_solution_4.py"
export_cell_content("solution_4", solution_file)
export_cell_content("tests_4", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    print('\033[91m your code presented some issue, check messages above', u'\u274C', \
        "\033[0m")
```

### Task 5

Create a Python dataclass called `CourseRegister`.
In the initialization method, the class should receive the `name` and `code` of the course, and save them as attributes.

The class should also have a method called `.register()` that receives an object of the class `Student` and adds it to the list of enrolled students.
- This list should be initialized as an empty list in the initialization method, and named `list_registered_students`.
- The objects should be added to the list only if they are an object of the `Student` class.
- The method should not add the same student object twice in the list.

The class should also have a method called `.remove()` that receives an object of the class `Student`, and removes it from the list of enrolled students. The method should check whether the student is present in the list, before trying to remove the student from the list.

**Suggestions:**
- You should run your test cells to tasks 2 and 4 for the test cell to work properly.
- To solve this problem, you need to use the knowledge from Chapter 9 (Object-oriented programming) of the lecture notes.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# solution cell
raise ValueError("Delete this line and write your solution here!")
```

```python
# uncomment the line(s) below to debug

c1 = CourseRegister('Applied object-oriented programming', 'EEN060')
print("New course created!")
print("List of enrolled students = ", c1.list_registered_students)
s1 = Student('John Doe', '201414202014', 'TIEPL')
print("Adding a student to the course...")
c1.register(s1)
print("List of enrolled students = ", c1.list_registered_students)
p1 = Person('Jack Dane', '201243202019', 'Manager')
print("Adding a person to the course...")
c1.register(p1)
print("List of enrolled students = ", c1.list_registered_students)
s2 = Student('Jane Doe', '201212202012', 'TIEPL', 'Gibraltargatan')
print("Adding a student to the course...")
c1.register(s2)
print("List of enrolled students = ", c1.list_registered_students)
print("Removing a student to the course...")
c1.remove(s1)
print("List of enrolled students = ", c1.list_registered_students)
print('Execution finished', u'\u2713')
```

```python
# test cell
from unittest.mock import patch


def test_solution_5() -> None:
    try:
        import solution_5
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_5 import CourseRegister
        from solution_4 import Student
        from solution_2 import Person
    except:
        raise ValueError("Your solution does not contain the right function!")
        
    assert 'CourseRegister' in dir(), 'Class CourseRegister does not exist!'
    assert 'register' in dir(CourseRegister), \
        'Class CourseRegister does not have method `register`'
    assert 'remove' in dir(CourseRegister), \
        'Class CourseRegister does not have method `remove`'

    c = CourseRegister('Applied object-oriented programming', 'EEN060')
    assert hasattr(c, 'name'), 'Class CourseRegister does not have property `name`.'
    assert hasattr(c, 'code'), 'Class CourseRegister does not have property `code`.'
    assert hasattr(c, 'list_registered_students'), \
        'Class CourseRegister does not have property `list_registered_students`.'
    assert isinstance(c.list_registered_students, list), \
        'The attribute list_enrolled_students is not a list'

    s1 = Student('John Doe', '201414202014', 'TIDSL')
    c.register(s1)
    assert len(c.list_registered_students) == 1

    p1 = Person('John Doe', '201414202014', 'Manager')
    c.register(p1)
    assert len(c.list_registered_students) == 1

    s2 = Student('Jane Doe', '201212202012', 'TKDES', 'Gibraltargatan')
    c.register(s2)
    assert len(c.list_registered_students) == 2

    c.remove(s1)
    assert len(c.list_registered_students) == 1

    c.remove(s1)
    assert len(c.list_registered_students) == 1

    c.remove(p1)
    assert len(c.list_registered_students) == 1

    c.register(s2)
    assert len(c.list_registered_students) == 1

    c = CourseRegister('Applied object-oriented programming', 'EEN060')
    assert isinstance(c.list_registered_students, list), \
        'The attribute list_enrolled_students is not a list'

    s1 = Student('John Doe', '201414202014_d', 'TIDSL')
    c.register(s1)
    assert len(c.list_registered_students) == 1

    p1 = Person('John Doe', '201414202014_d', 'Manager')
    c.register(p1)
    assert len(c.list_registered_students) == 1

    s2 = Student('Jane Doe', '201212202012_d', 'TKDES', 'Gibraltargatan')
    c.register(s2)
    assert len(c.list_registered_students) == 2

    c.remove(s1)
    assert len(c.list_registered_students) == 1

    c.remove(s1)
    assert len(c.list_registered_students) == 1

    c.remove(p1)
    assert len(c.list_registered_students) == 1
```

```python
# validate the solution
solution_file = "solution_5.py"
test_file = "test_solution_5.py"
export_cell_content("solution_5", solution_file)
export_cell_content("tests_5", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    print('\033[91m your code presented some issue, check messages above', u'\u274C', \
        "\033[0m")
```

### Task 6

Create a Python class named `Word`.

The class must have:

- An attribute `text: str`
- An initialization method `__init__(text: str)` that saves `text` as an attribute

The class must implement the magic method: `__eq__` that takes another `Word` object as input and returns a `bool`. 

Two `Word` objects are considered equal if their `text` values are equal **ignoring case**.

Example:

```python
Word("ha") == Word("HA")   # True
Word("ha") == Word("eh")   # False

```python
# solution cell
raise ValueError("Delete this line and write your solution here!")
```

```python
# uncomment the line(s) below to debug

word1 = Word('HA')
word2 = Word('ha')
word3 = Word('eh')
print(f"word1 has same text as word2 (ignoring case): {word1 == word2}")
print(f"word2 has different text as word3: {word2 != word3}")
print('Execution finished', u'\u2713')
```

```python
# test cell
def test_solution_6() -> None:
    try:
        import solution_6
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_6 import Word
    except:
        raise ValueError("The module does not have the necessary function!")

    w = Word("ha")
    assert hasattr(w, "text"), "Word must have an attribute `text`."
    assert w.text == "ha", "`text` must store the given input string."

    assert Word("ha") == Word("HA"), "Equality must ignore case."
    assert Word("ha") != Word("eh"), "Different words must not be equal."

    assert (Word("ha") == Word("ha")) is True
```

```python
# validate the solution
solution_file = "solution_6.py"
test_file = "test_solution_6.py"
export_cell_content("solution_6", solution_file)
export_cell_content("tests_6", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    print('\033[91m your code presented some issue, check messages above', u'\u274C', \
        "\033[0m")
```