# EEN065 - Applied object-oriented programming


Teacher: [Seyed Morteza Ahmadian](https://www.chalmers.se/en/persons/seyedah/) and [Carlos Natalino](https://www.chalmers.se/en/persons/carda/) / Examiner: [Paolo Monti](https://www.chalmers.se/en/persons/mpaolo/)

[Canvas course page](https://chalmers.instructure.com/courses/39015)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before submitting this assignment, please ensure that everything runs correctly.

First, **restart the kernel**, then **run all cells**.
Next, verify that all tests execute successfully.

Note that if any problem produces an error, the subsequent problems **will not** be tested.

In case of discrepancies between the problem description and the tests, you should resolve the issue with the tests in mind, as they determine correctness.

There are four types of cells in this notebook:
1.	*Solution cells*: These are the cells where you write your answers or modify the existing code to solve the problem.
2.	*Debug cells*: These cells are for debugging your solution. You may add or edit code here, but ensure that it does not introduce any errors.
3. *Test cells:*  These cells run **automated tests** to verify whether your solution is **functionally correct**.
4. *Code quality tests:*  These cells check whether your solution follows the **code quality standards** of the course. Even if your code produces the correct output, it must also be **clean, readable, and well-structured**.

- If all tests pass, you will see an output similar to:
  - `1 passed`

- If a test fails, you will see an **error message** explaining:
  - Which input was tested
  - What output was expected
  - What your function actually returned

You will see several sections in the output:

- **tests:** Runs the test cells with coverage enabled.  
  - `1 passed` indicates success.
- **coverage:** Checks for any line of your solution that is not useful.  
  - `100% coverage` indicates success.
- **flake8:** Checks for style issues (e.g., unused variables, line length).  
  - Any output other than `0` here means you must fix style violations.
- **pylint:** Performs deeper code analysis and assigns a score.  
  - A score of `10.00/10` means your code fully complies.
- **mypy**: Checks that your code respects the specified type annotations.
  - Any reported error means a type hint is missing or incorrect.
- **isort:** Verifies that imports are correctly ordered.  
  - If differences are shown, reorder your imports.

At the end:
- **`tests passed ✓`** → Your solution is correct *and* meets all quality requirements.
- **`your code presented some issue`** → Review the messages above, fix the reported issues, and rerun the cell.

**You must pass both functional tests and code quality checks to receive full credit.**

Common errors and what they mean:
- **`You did not execute your solution cell!`**  
  → You forgot to run the solution cell before running the test cell.
- **`The module does not have the necessary function!`**  
  → The required function is missing or named incorrectly.
- **AssertionError**  
  → Your function ran, but it returned an incorrect result for at least one test case.

**Fix all test errors before moving on**, as later problems will not be tested if an earlier one fails.

**Delete** the line `raise NotImplementedError()` from any problem you solve.

**Do not add or remove any cells in this file.** All required cells are already provided.

To execute a cell, select it and press **CTRL+Enter** (Windows), **CMD+Enter** (macOS), or click the **Run cell** button.

## Assignment Week 4

**Observations:**
For the problems in this assignment, you should:
- make sure that your `utils.py` file has the following line.
- you will find the set of references needed to solve this assignment in the assignment page in canvas.

**Preparation:** Run the cell below every time you start working on this file, and every time you restart the kernel.

```python
%load_ext autoreload
import os

# getting latest version of utils.py
# import requests
# response = requests.get("https://raw.githubusercontent.com/carlosnatalino/programming-assignments/refs/heads/main/utils.py")
# with open("utils.py", "wb") as f:
#     f.write(response.content)

if not os.path.exists("utils.py"):
    raise ValueError("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.")

# check if this file is within the same folder as the `utils.py` file.
try:
    from utils import export_cell_content
except:
    import sys
    raise ValueError("Did you select the right kernel that starts with venv?")

filename = "assignment_week_04_public.ipynb"
assert filename != None, "The file name must be set."
if not os.path.exists(filename):
    raise ValueError(f"The value of the `filename` variable is wrong. \
The file `{filename}` does not exist. Write the correct name of this file!")

os.environ["NOTEBOOK_NAME"] = os.path.join(os.getcwd(), filename)
```

### Task 1

Create a function called `assignment_points` that receives a list containing __seven__ string items as parameter.
The value of each item is either `Approved` or `Not Approved`.
Each value in the list represents whether a particular assignment is approved or not.
The function then returns a float that represents the number of assignment points obtained according to the scale defined in the [course PM](https://chalmers.instructure.com/courses/38213#:~:text=1%20AP,15%20APs).

Make sure to return `0` if no assignment is approved, i.e., when you recieve `["Not Approved","Not Approved","Not Approved","Not Approved","Not Approved","Not Approved","Not Approved"]`.
Also return `0` if the number of items in the received list is not 7.

**Note**: This task is the same as Problem 1 in the last assignment, with the only difference being the input format. Whereas in last assignment you were given the number of approved assignment, now you need to calculate this number based on the input list.

# receive a list of assignments
# check if length of list is not 7 return 0
# count how many items are "Approved"
# create a list of points
# get the value from the list at position count
# return the value

```python
"""Solution for assignment points calculation."""


def assignment_points(assignments_list: list[str]) -> int:
    """Return points based on number of 'Approved' assignments."""
    if len(assignments_list) != 7:
        return 0
    count = assignments_list.count("Approved")
    points = [0, 1, 2, 4, 6, 9, 12, 15]
    return points[count]
```

```python
### uncomment the line(s) below to debug

assignments = ["Approved","Approved","Approved","Approved","Approved", \
            "Approved","Approved"]
ap = assignment_points(assignments)
print(f"The assignment points are {ap}")
print("Execution finished", "\u2713")
```

Output:
```
The assignment points are 15
Execution finished ✓
```

```python
# test cell
def test_solution_1() -> None:
    try:
        import solution_1
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_1 import assignment_points
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [(["Not Approved","Not Approved","Not Approved","Not Approved", \
        "Not Approved","Not Approved","Not Approved"], 0.0), \
        (["Approved","Not Approved","Not Approved","Not Approved", "Not Approved", \
            "Not Approved","Not Approved"], 1.0), \
        (["Approved","Approved","Not Approved","Not Approved", "Not Approved", \
            "Not Approved","Not Approved"], 2.0), \
        (["Approved","Approved","Approved","Not Approved", "Not Approved", \
            "Not Approved","Not Approved"], 4.0), \
        (["Approved","Approved","Approved","Approved", "Not Approved", "Not Approved", \
            "Not Approved"], 6.0), \
        (["Approved","Approved","Approved","Approved", "Approved","Not Approved", \
            "Not Approved"], 9.0), \
        (["Approved","Approved","Approved","Approved", "Approved","Approved", \
            "Not Approved"], 12.0), \
        (["Approved","Approved","Approved","Approved", "Approved","Approved", \
            "Approved"], 15.0), \
        (["Approved","Approved","Approved","Approved", "Approved","Approved", \
            "Approved", "Not Approved"], 0.0), \
        (["Approved","Approved","Approved","Approved", "Approved","Approved", \
            "Approved", "Approved"], 0.0)]

    for _in, _out in test_cases:
        _res = assignment_points(_in)
        assert _res == _out, f"The function received the input\n`{_in}` \
        and should return the value `{_out}`\n but returned the value `{_res}`."
```

```python
# validate the solution
solution_file = "solution_1.py"
test_file = "test_solution_1.py"
export_cell_content("solution_1", solution_file)
export_cell_content("tests_1", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-28 13:11:39.269085 (3 minutes ago)
Latest saved date of the file: 2026-04-28 13:11:39.269085 (3 minutes ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.09s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss  Cover   Missing
---------------------------------------------
solution_1.py       6      0   100%
---------------------------------------------
TOTAL               6      0   100%
coverage True

[0m
############################################## flake8 ##############################################
flake8 True

[0m
############################################## pylint ##############################################

----------------------------------------
```

### Task 2

Create a Python function named `reverse_lower_and_upper_case` that receives a list of strings as a parameter.
The function must iterate over each element in the list and reverse the case of string.
That is, if an element is in upper case, then convert it to lower case.
If it is in lower case, then convert it to upper case.

**Note**: This task is the same as Problem 3 in the last assignment, with the only difference being the input format. Whereas in last assignment you were asked to change the case of only one string, now you need to change the case of all the strings in a list.

**Examples**

Input: ['test', 'PRACTICE']

Output: ['TEST', 'practice']

Explanation: Here 'test' is in lower case, so we converted it to upper case and 'PRACTICE' is in upper case so we converted it into lower case.

# receive a list of strings
# create an empty list called result
# loop through each item in the list
# check if item is lowercase
# convert it to uppercase and add to result
# otherwise if item is uppercase
# convert it to lowercase and add to result
# return the result list

```python
"""Module for reversing case of strings in a list."""


def reverse_lower_and_upper_case(items: list[str]) -> list[str]:
    """Reverse lowercase and uppercase strings in a list."""
    result = []
    for item in items:
        if item.islower():
            result.append(item.upper())
        elif item.isupper():
            result.append(item.lower())
    return result
```

```python
### uncomment the line(s) below to debug

initial = ['test', 'PRACTICE']
result = reverse_lower_and_upper_case(initial)
print(f"The reverse case version of {initial} is {result}")
print("Execution finished", "\u2713")
```

Output:
```
The reverse case version of ['test', 'PRACTICE'] is ['TEST', 'practice']
Execution finished ✓
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
# test cell
def test_solution_2() -> None:
    try:
        import solution_2
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_2 import reverse_lower_and_upper_case
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [(['test', 'PRACTICE'], ['TEST', 'practice']), \
        (['TEST', 'practice'], ['test', 'PRACTICE']), \
        (['TEST', 'PRACTICE'], ['test', 'practice']), \
        (['resource granted', 'RESOURCE NOT GRANTED', 'grattis'], \
        ['RESOURCE GRANTED', 'resource not granted', 'GRATTIS'])]

    for _in, _out in test_cases:
        _res = reverse_lower_and_upper_case(_in)
        assert _res == _out, f"The function received the input\n`{_in}` \
        and should return the value `{_out}`\n but returned the value `{_res}`."
```

```python
# validate the solution
solution_file = "solution_2.py"
test_file = "test_solution_2.py"
export_cell_content("solution_2", solution_file)
export_cell_content("tests_2", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-28 13:28:01.575059 (2 seconds ago)
Latest saved date of the file: 2026-04-28 13:28:01.575059 (2 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.03s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss  Cover   Missingcoverage True

[0m
############################################## flake8 ##############################################

---------------------------------------------
solution_2.py       8      0   100%
---------------------------------------------
TOTAL               8      0   100%
flake8 True

[0m
############################################## pylint ##############################################
pylint True

[0m
#######################
```

### Task 3

Create a Python function named `age_counter`.
It receives __one string__ as parameter.
The string contains name and age of people, separated by commas as shown below: 

`"James = 20, Rock = 32, Fred = 20"`

The function must return a dictionary where each unique age becomes the key (integer), and the value associated with the key is the number of people with that given age in the input. The function must also print an integer representing the number of people within the age 20 to 30 (both inclusive).

**Note:**
- The equal sign, commas, and spaces must not be considered (i.e., must be removed from the text).
- Knowing in which stage to remove the equal sign, commas, and spaces, it one of the knowledge to be acquired in this task.
- You must not use external functions to solve this problem, i.e., you must not use the `import` command.

# receive a string with names and ages
# remove all spaces from the string
# split the string into a list using comma
# create an empty dictionary called ages
# create a counter set to 0
# for each person in the list
# extract the age from the string
# check if age is already in dictionary
# increase its count by 1
# otherwise add age to dictionary with value 1
# check if age is between 20 and 30
# increase counter by 1
# print the counter
# return the dictionary

```python
"""Solution for task 3."""


def age_counter(text):
    """Count how many people have each age and print ages from 20 to 30."""
    text = text.replace(" ", "")
    people = text.split(",")
    ages = {}
    counter = 0

    for person in people:
        age = int(person.split("=")[1])

        if age in ages:
            ages[age] += 1
        else:
            ages[age] = 1

        if 20 <= age <= 30:
            counter += 1

    print(counter)
    return ages
```

```python
### uncomment the line(s) below to debug

text = "Alice = 25, Bob = 31, Charlie = 22, David = 22, Emma = 28, \
        Frank = 35, Grace = 25, Harry = 25, Isabella = 22, Jack = 27"
print("Number of people within age 20 to 30 is")
age_count = age_counter(text)
print("Age statistics are:\n",str(dict(age_count)).replace(', ',',\n '))
print("Execution finished", "\u2713")
```

Output:
```
Number of people within age 20 to 30 is
8
Age statistics are:
 {25: 3,
 31: 1,
 22: 3,
 28: 1,
 35: 1,
 27: 1}
Execution finished ✓
```

```python
# test cell
from unittest.mock import patch
from io import StringIO

def test_solution_3() -> None:
    try:
        import solution_3
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_3 import age_counter
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [("Alice = 25, Bob = 31, Charlie = 22, David = 22, Emma = 28, \
        Frank = 35, Grace = 25, Harry = 25, Isabella = 22, Jack = 27", \
            {25: 3, 31: 1, 22: 3, 28: 1, 35: 1, 27: 1}, "8"), \
        ("Sophia = 24, Olivia = 28, Liam = 32, Noah = 24, \
            Emily = 28, Ava = 24, Benjamin = 35, Mia = 31, Elijah = 28, Harper = 28", \
                {24: 3, 28: 4, 32: 1, 35: 1, 31: 1}, "7")]
    for _in, _out, _print in test_cases:
        with patch("sys.stdout", new=StringIO()) as fakeOutput:
            _res = age_counter(_in)
            assert _res == _out, f"The function received input\n`{_in}`\
                and should return the value `{_out}` of type \
    `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
            assert fakeOutput.getvalue().strip() == _print, f"For the received input \
    received input\n`{_in}`\nthe function should print \
    the value `{_print}` but printed the value `{fakeOutput.getvalue().strip()}`."
```

```python
# validate the solution
solution_file = "solution_3.py"
test_file = "test_solution_3.py"
export_cell_content("solution_3", solution_file)
export_cell_content("tests_3", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-28 13:32:18.679211 (2 seconds ago)
Latest saved date of the file: 2026-04-28 13:32:18.679211 (2 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.08s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss  Cover   Missing
---------------------------------------------
solution_3.py      14      0   100%
---------------------------------------------
TOTAL              14      0   100%
coverage True

[0m
############################################## flake8 ##############################################
flake8 True

[0m
############################################## pylint ##############################################

----------------------------------------
```

### Task 4

Create a Python function called `numbers_and_bools` that receives a list as parameter.
The list may contain numbers (integers or floats), or Boolean values, or strings.

The function must return two lists.
The first list must contain all the Boolean values.
The second list must contain all the numbers (integers or floats) from the received list. 
The strings must be ignored, i.e., must not be included in the list.

If there are no Boolean values or numbers in the list, the function must return two empty lists, accordingly.

# receive a list
# create an empty list called bools
# create an empty list called nums
# for each item in the list
# check if item is a boolean
# add it to bools list
# otherwise if item is a number (int or float)
# add it to nums list
# ignore strings
# return both lists (bools and nums)

```python
"""Solution for task 4."""


def numbers_and_bools(lst):
    """Return booleans and numbers from the list."""
    bools = []
    nums = []

    for item in lst:
        if isinstance(item, bool):
            bools.append(item)
        elif isinstance(item, (float, int)):
            nums.append(item)

    return bools, nums
```

```python
### uncomment the line(s) below to debug

data_structure = [True, 10., False, 90., True, 45., True, 12, -1.2, 44.9]
# data_structure = ['this', 10., 'is', 90., 'my', 45., 'Python', 12, 'code', 44.9]
result = numbers_and_bools(data_structure)
print(f"The set of values: {data_structure}")
print(f"Results in bools: {result[0]}")
print(f"and numbers: {result[1]}")
print("Execution finished", u"\u2713")
```

Output:
```
The set of values: [True, 10.0, False, 90.0, True, 45.0, True, 12, -1.2, 44.9]
Results in bools: [True, False, True, True]
and numbers: [10.0, 90.0, 45.0, 12, -1.2, 44.9]
Execution finished ✓
```

```python
# test cell
def test_solution_4() -> None:
    try:
        import solution_4
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_4 import numbers_and_bools
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [([True, 10., False, 90., True, 45., True, 12, -1.2, 44.9], \
        ([True, False, True, True], [10.0, 90.0, 45.0, 12, -1.2, 44.9])), \
            (['this', 10., 'is', 90., 'another', 45., 'case', 12, 'test', 44.9], \
        ([], [10.0, 90.0, 45.0, 12, 44.9]))]
    for _in, _out in test_cases:
        _res = numbers_and_bools(_in)
        assert _res == _out, f"The function received the input\n`{_in}` \
        and should return the value `{_out}`\n but returned the value `{_res}`."
```

```python
# validate the solution
solution_file = "solution_4.py"
test_file = "test_solution_4.py"
export_cell_content("solution_4", solution_file)
export_cell_content("tests_4", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-28 13:58:21.894196 (3 seconds ago)
Latest saved date of the file: 2026-04-28 13:58:21.894196 (3 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.03s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss  Cover   Missing
---------------------------------------------
solution_4.py       9      0   100%
---------------------------------------------
TOTAL               9      0   100%
coverage True

[0m
############################################## flake8 ##############################################
flake8 True

[0m
############################################## pylint ##############################################

----------------------------------------
```

### Task 5

Create a Python function named `lists_intersection` that receives two lists as an input.
The function should make sure that the inputs are of the `list` type.
If either of the inputs is not a list, the function should return `None` (the *void* type in Python, not a string with `"None"` inside).

Otherwise, the function should return a list that contains all the elements that are common to both the lists. If there are no such elements, return an empty list `[]`.

# receive two inputs (a and b)
# check if both inputs are lists
# if not return None
# create an empty list called result
# for each item in list a
# check if item exists in list b
# and item is not already in result
# add item to result
# return the result list

```python
"""Solution for task 5."""


def lists_intersection(a, b):
    """Return common elements."""
    if not isinstance(a, list) or not isinstance(b, list):
        return None

    result = []

    for item in a:
        if item in b and item not in result:
            result.append(item)

    return result
```

```python
# uncomment the line(s) below to debug

list_1 = ['a', 'w', 'b', 'x']
list_2 = ['a', 'b', 'c']
result = lists_intersection(list_1, list_2)
print(f"The intersection of {list_1} and {list_2} is {result}")
print("Execution finished", "\u2713")
```

Output:
```
The intersection of ['a', 'w', 'b', 'x'] and ['a', 'b', 'c'] is ['a', 'b']
Execution finished ✓
```

```python
# test cell
from unittest.mock import patch


def test_solution_5() -> None:
    try:
        import solution_5
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_5 import lists_intersection
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [((('test_1', 'test_2'), ('test_3', 'test_4')), None), \
        ((('test_1', 'test_2'), ['test_3', 'test_4']), None), \
            ((['a', 'b'], ['b', 'c']), ['b']), \
        ((['a', 'b'], ['a', 'b']), ['a', 'b']), ((['a', 'b'], ['c', 'd']), [])]

    for _in, _out in test_cases:
        _res = lists_intersection(*_in)
        if _out:
            assert set(_res) == set(_out), f"The function received the input\n`{_in}` \
            and should return the value `{_out}`\n but returned the value `{_res}`."
        elif _out == []:
            assert _out == _res, f"The function received the input\n`{_in}` \
            and should return the value `{_out}`\n but returned the value `{_res}`."            
        else:
            assert _res == None, f"The function received the input\n`{_in}` \
            and should return the value `{_out}`\n but returned the value `{_res}`."

    with patch('builtins.set') as mock_set:
        lists_intersection(['a', 'b'], ['c', 'd'])
    mock_set.assert_not_called()
```

```python
# validate the solution
solution_file = "solution_5.py"
test_file = "test_solution_5.py"
export_cell_content("solution_5", solution_file)
export_cell_content("tests_5", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-28 13:58:21.894196 (5 seconds ago)
Latest saved date of the file: 2026-04-28 13:58:21.894196 (5 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[0m
############################################# coverage #############################################

[32m[32m[1m1 passed[0m[32m in 0.08s[0m[0m
Name            Stmts   Miss  Cover   Missing
---------------------------------------------
solution_5.py       8      0   100%
---------------------------------------------
TOTAL               8      0   100%
coverage True

[0m
############################################## flake8 ##############################################
flake8 True

[0m
############################################## pylint ##############################################

----------------------------------------
```

### Task 6

Create a Python function called `person_number_validate` that receives a 10 digit number as a string and validates whether the number provided is correct person number (returning `True`) or not (returning `False`). 

Rules for person number:
- first six digits must represent a valid date in YYMMDD format. You may ignore the leap year condition
- the last digit should be a valid checksum. More information here: https://en.wikipedia.org/wiki/Personal_identity_number_(Sweden)#Checksum


```python
"""Solution for task 6."""


def person_number_validate(pnr):
    """Validate Swedish personal number."""
    if not isinstance(pnr, str) or len(pnr) != 10 or not pnr.isdigit():
        return False

    month = int(pnr[2:4])
    day = int(pnr[4:6])

    days_in_month = [31, 28, 31, 30, 31, 30, 31, 31, 30, 31, 30, 31]

    if month < 1 or month > 12:
        return False

    if day < 1 or day > days_in_month[month - 1]:
        return False

    total = 0

    for index in range(9):
        number = int(pnr[index])

        if index % 2 == 0:
            number *= 2

        if number > 9:
            number -= 9

        total += number

    check_digit = (10 - total % 10) % 10

    return check_digit == int(pnr[9])
```

```python
# uncomment the line(s) below to debug

# number = "8112289874" # Valid
number = "3101057853" # Valid
valid = person_number_validate(number)
valid_string = "valid" if valid else "invalid"
print(f"The person number {number} is {valid_string}")
print("Execution finished", "\u2713")
```

Output:
```
The person number 3101057853 is valid
Execution finished ✓
```

```python
# test cell
from datetime import datetime


def test_solution_6() -> None:
    try:
        import solution_6
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_6 import person_number_validate
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [
        # Valid
        ('8112289874', True),
        ('3101057853', True),

        # Invalid checksum
        ('3101057851', False),

        # Invalid month
        ('9913211234', False),   # month 13
        ('1200301234', False),   # month 00

        # Invalid day for 30-day month
        ('1211310123', False),   # Nov 31

        # Invalid February dates
        ('1202301212', False),   # Feb 30
        ('0002301234', False),   # Feb 30

        # Clearly invalid day
        ('1200321234', False),   # day 32
    ]

    for _in, _out in test_cases:
        _res = person_number_validate(_in)
        assert _res == _out, f"The function received the input\n`{_in}` \
        and should return the value `{_out}`\n but returned the value `{_res}`."
```

```python
# validate the solution
solution_file = "solution_6.py"
test_file = "test_solution_6.py"
export_cell_content("solution_6", solution_file)
export_cell_content("tests_6", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

# print('\n\033[0m')
# print(" coverage ".center(100, "#"))
# !coverage report -m {solution_file}
# if _exit_code != 0:
#     passed = False
# print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-28 13:58:21.894196 (8 seconds ago)
Latest saved date of the file: 2026-04-28 13:58:21.894196 (8 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.04s[0m[0m

[0m
############################################## flake8 ##############################################
flake8 True

[0m
############################################## pylint ##############################################

--------------------------------------------------------------------
Your code has been rated at 10.00/10 (previous run: 10.00/10, +0.00)

pylint True

[0m
############################################### mypy ###############################################
[1m[92mSuccess: no issues found in 1 source file[0mmypy True

[0m
###########################################
```

### Task 7

Create a Python function called `validating_sum_of_squares` that receives a sequence of elements and validates whether or not each of the elements is a number and a sum of squares (eg. 25 is sum of 16 and 9 that is $4^2 + 3^2$).

Note that numbers can be placed within a string. And, for any integer $a$, since $a^2 = 0^2 + a^2$, square of any integer is assumed to valid.

At the end, the function should return a sequence (with the same number of elements as the input) containing `True` for the elements that are sum of squares and `False` otherwise.

**Note:** The input sequence may contain elements that are not numbers and cannot be converted to numbers. For example, a boolean value `None` and a string value`"asd"`. These elements should be assigned False in the returned list. To identify numbers from invalid inputs, you should use the `try`/`except` blocks from section 9.4, and the convertion to `float` from section 3.4.3 in the lecture notes.

**Examples**

Input: ['5', 3, 25, None, 'asd']

Output: [True, False, True, False, False]

```python
"""Solution for task 7."""


def validating_sum_of_squares(sequence):
    """Return True if element is sum of two squares, else False."""
    result = []

    for item in sequence:
        try:
            number = float(item)
        except (TypeError, ValueError):
            result.append(False)
            continue

        # måste vara ett icke-negativt heltal
        if number < 0 or number != int(number):
            result.append(False)
            continue

        number = int(number)
        valid = False

        # smartare loop (behöver inte gå hela vägen till number)
        for i in range(int(number ** 0.5) + 1):
            remaining = number - i * i
            j = int(remaining ** 0.5)

            if j * j == remaining:
                valid = True
                break

        result.append(valid)

    return result
```

```python
# uncomment the line(s) below to debug

seq_input = ['5', 3, 25, None, 'asd']
x = validating_sum_of_squares(seq_input)
print(f' Input sequence: ', '\t'.join([str(x) for x in seq_input]))
print(f'Output sequence: ', '\t'.join([str(r) for r in x]))
print('Execution finished', u'\u2713')
```

Output:
```
Input sequence:  5	3	25	None	asd
Output sequence:  True	False	True	False	False
Execution finished ✓
```

```python
# test cell
from unittest.mock import patch


def test_solution_7() -> None:
    try:
        import solution_7
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_7 import validating_sum_of_squares
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [(['2', 'x', -10, 3.3, 'asd', None, 'b', 4.0], \
        [True, False, False, False, False, False, False, True]), \
        ([9, 'x', -9, None], [True, False, False, False])]

    for _in, _out in test_cases:
        _res = validating_sum_of_squares(_in)
        assert _res == _out, f"The function received the input\n`{_in}` \
        and should return the value `{_out}`\n but returned the value `{_res}`."
```

```python
# validate the solution
solution_file = "solution_7.py"
test_file = "test_solution_7.py"
export_cell_content("solution_7", solution_file)
export_cell_content("tests_7", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

# print('\n\033[0m')
# print(" coverage ".center(100, "#"))
# !coverage report -m {solution_file}
# if _exit_code != 0:
#     passed = False
# print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-28 13:58:21.894196 (10 seconds ago)
Latest saved date of the file: 2026-04-28 13:58:21.894196 (10 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[0m
############################################## flake8 ##############################################

[32m[32m[1m1 passed[0m[32m in 0.07s[0m[0m
flake8 True

[0m
############################################## pylint ##############################################

--------------------------------------------------------------------
Your code has been rated at 10.00/10 (previous run: 10.00/10, +0.00)

pylint True

[0m
############################################### mypy ###############################################
[1m[92mSuccess: no issues found in 1 source file[0m
mypy True

[0m
########################################
```